rule_files:
  - alert.rules.yml
tests:
  - interval: 1m
    input_series:
      - series: 'gossip_convergence_seconds_bucket{le="30"}'
        values: '0+10x31'
      - series: 'gossip_convergence_seconds_bucket{le="60"}'
        values: '0+20x31'
      - series: 'gossip_convergence_seconds_bucket{le="+Inf"}'
        values: '0+20x31'
      - series: 'CONSUMER_FEE_P90'
        values: '1000000+0x31'
      - series: 'param_change_active{key="ConsumerFeeComfortP90Microunits"}'
        values: '500000+0x31'
      - series: 'INDUSTRIAL_DEFERRED_TOTAL'
        values: '0+4x31'
      - series: 'INDUSTRIAL_ADMITTED_TOTAL'
        values: '0+6x31'
      - series: 'subsidy_cpu_ms_total'
        values: '0+1x31'
      - series: 'subsidy_bytes_total'
        values: '0+1x31'
      - series: 'rent_escrow_locked_total'
        values: '0+2x31'
      - series: 'bridge_metric_delta{metric="bridge_reward_claims_total",peer="bridge-node",labels=""}'
        values: '2+0x24 40+0x7'
      - series: 'bridge_metric_rate_per_second{metric="bridge_reward_claims_total",peer="bridge-node",labels=""}'
        values: '0.3+0x24 5+0x7'
      - series: 'bridge_metric_delta{metric="bridge_reward_claims_total",peer="bridge-node",labels="asset=btc"}'
        values: '0.5+0x24 9+0x7'
      - series: 'bridge_metric_rate_per_second{metric="bridge_reward_claims_total",peer="bridge-node",labels="asset=btc"}'
        values: '0.05+0x24 0.4+0x7'
      - series: 'bridge_metric_delta{metric="bridge_reward_claims_total",peer="bridge-node",labels="asset=eth"}'
        values: '0.5+0x31'
      - series: 'bridge_metric_rate_per_second{metric="bridge_reward_claims_total",peer="bridge-node",labels="asset=eth"}'
        values: '0.05+0x31'
      - series: 'bridge_remediation_dispatch_ack_total{action="escalate",playbook="governance-escalation",target="http",state="pending"}'
        values: '0+1x31'
      - series: 'bridge_remediation_dispatch_total{action="escalate",playbook="governance-escalation",target="http",status="success"}'
        values: '0+1x31'
      - series: 'bridge_remediation_dispatch_ack_total{action="escalate",playbook="governance-escalation",target="http",state="closed"}'
        values: '0+0x31'
      - series: 'range_boost_forwarder_fail_total'
        values: '0+1x31'
      - series: 'range_boost_toggle_latency_seconds_count'
        values: '0+0x31'
      - series: 'treasury_executor_last_submitted_nonce'
        values: '0+1x31'
      - series: 'treasury_executor_lease_last_nonce'
        values: '5+0x16 5+0x7 2+0x8'
      - series: 'treasury_executor_lease_released'
        values: '0+0x25 1+0x6'
    alert_rule_test:
      - eval_time: 20m
        alertname: ConvergenceLag
        exp_alerts:
          - exp_labels:
              severity: page
            exp_annotations:
              summary: "Chain convergence p95 above 30s"
      - eval_time: 20m
        alertname: ConsumerFeeComfortBreached
        exp_alerts:
          - exp_labels:
              severity: warn
            exp_annotations:
              summary: "Consumer fee p90 above comfort threshold"
      - eval_time: 30m
        alertname: IndustrialDeferralRatioHigh
        exp_alerts:
          - exp_labels:
              severity: warn
            exp_annotations:
              summary: "Industrial deferral ratio above 30%"
      - eval_time: 30m
        alertname: SubsidyGrowthSpike
        exp_alerts:
          - exp_labels:
              severity: warn
            exp_annotations:
              summary: "Subsidy counters growing unusually fast"
      - eval_time: 30m
        alertname: RentEscrowLockedSpike
        exp_alerts:
          - exp_labels:
              severity: warn
            exp_annotations:
              summary: "Rent escrow locked BLOCK spiking"
      - eval_time: 30m
        alertname: BridgeCounterDeltaSkew
        exp_alerts:
          - exp_labels:
              severity: warn
              metric: bridge_reward_claims_total
              peer: bridge-node
              labels: ""
            exp_annotations:
              summary: "Bridge counter delta skew detected for bridge_reward_claims_total on bridge-node ()"
      - eval_time: 30m
        alertname: BridgeCounterRateSkew
        exp_alerts:
          - exp_labels:
              severity: warn
              metric: bridge_reward_claims_total
              peer: bridge-node
              labels: ""
            exp_annotations:
              summary: "Bridge counter rate skew detected for bridge_reward_claims_total on bridge-node ()"
      - eval_time: 30m
        alertname: BridgeCounterDeltaLabelSkew
        exp_alerts:
          - exp_labels:
              severity: warn
              metric: bridge_reward_claims_total
              peer: bridge-node
              labels: "asset=btc"
            exp_annotations:
              summary: "Bridge counter delta skew detected for bridge_reward_claims_total on bridge-node (asset=btc)"
      - eval_time: 30m
        alertname: BridgeCounterRateLabelSkew
        exp_alerts:
          - exp_labels:
              severity: warn
              metric: bridge_reward_claims_total
              peer: bridge-node
              labels: "asset=btc"
            exp_annotations:
              summary: "Bridge counter rate skew detected for bridge_reward_claims_total on bridge-node (asset=btc)"
      - eval_time: 30m
        alertname: BridgeRemediationAckPending
        exp_alerts:
          - exp_labels:
              severity: warn
              action: escalate
              playbook: governance-escalation
              target: http
              state: pending
            exp_annotations:
              summary: "Bridge remediation acknowledgements stalled (pending) for escalate on target http"
      - eval_time: 30m
        alertname: BridgeRemediationClosureMissing
        exp_alerts:
          - exp_labels:
              severity: warn
              action: escalate
              playbook: governance-escalation
              target: http
            exp_annotations:
              summary: "Bridge remediation closures missing for escalate on target http"
      - eval_time: 20m
        alertname: RangeBoostForwarderFailures
        exp_alerts:
          - exp_labels:
              severity: page
            exp_annotations:
              summary: "RangeBoost forwarder failures accumulating without recent toggles"
      - eval_time: 22m
        alertname: TreasuryLeaseWatermarkLagging
        exp_alerts:
          - exp_labels:
              severity: page
            exp_annotations:
              summary: "Treasury executor lease watermark stuck behind submitted nonces"
      - eval_time: 30m
        alertname: TreasuryLeaseWatermarkRegression
        exp_alerts:
          - exp_labels:
              severity: page
            exp_annotations:
              summary: "Treasury executor lease watermark regressed"
      - eval_time: 30m
        alertname: TreasuryLeaseReleased
        exp_alerts:
          - exp_labels:
              severity: warn
            exp_annotations:
              summary: "Treasury executor lease released"
